---
title: 'Type Hinting'
description: 'In this chapter, you will learn how to add type hints to your variables and function definitions so that you can perform static checking on your scripts. This will help ensure that your code is tested and will not produce surprising behavior in production.'
---

## Adding Type Hints to Function Definitions

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Adding Type Hints to Function Definitions

```yaml
type: VideoExercise
key: 9219a74cd2
xp: 50
```

`@projector_key`
837d120fc1f78e8601d7a6765e6a5e43
